```python
# Initialize Otter
import otter
grader = otter.Notebook("ps0.ipynb")
```

# STATS 507 W22
## Problem Set 0

This problem set is designed to introduce to Jupyter notebooks and our autograder platform. It is easier and shorter than the rest of the problem sets.

The first cell imports the autograder package that we are using to autograde your assignments. If running the above cell generates an error, it means that you do not have this package installed. Install it by uncommenting and running the following command:

```python
# %pip install otter-grader
```

You only need to install this package once. Restart your kernel after installing and verify that the `otter` package loads. 

## Autograded questions
Questions 1-4 are autograded. You should fill in the answers where indicated, and then submit this notebook to Gradescope for grading as detailed in the assignment instructions.

### Question 1
Print "Hello world!!" to the console.

```python
print("Hello world!!")
```

Output:
```
Hello world!!
```

### Question 2
Assign the variable `x` to be the smallest prime number.

```python
x = 2
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### Question 3
Write a function `howlong(s)` that takes a string argument `s` and returns the length of that string.

```python
def howlong(s):
    return len(s)
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Question 4
Assign to the variable `nmin_2024` the number of minutes in the year 2024.

```python
# 2024 is a leap year
nmin_2024 = 366 * 24 * 60
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Manually graded
Questions 5 and 6 are manually graded. For PS1 onwards, the graders would look at these responses and grade them by hand. For this problem set, we will use them to collect some basic information about your computing background and interests.

<!-- BEGIN QUESTION -->

### Question 5
How much programming experience do you have?

I have self-learned python last year and use python in IOE610&IOE512. I am familiar with dynamic programming and linear programming especially coding via gurobi.

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 6
What are you hoping to get out of this course?

I hope to have a systematic knowledge about python and in the future I am considering to learn some machine learning and reforcement learning. Hopefully this course will be a solid foundation for my future work.

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q2 (points: 1)
```
>>> isinstance(x, int)
True
```
```
>>> x == 2
True
```


Tests for q3 (points: None)
```
>>> howlong("") == 0
True
```


Tests for q4 (points: None)
