```python
# Initialize Otter
import otter
grader = otter.Notebook("ps7.ipynb")

```

# PS7: Regular expressions
In this problem set set you will get some basic practice using regular expressions. In Python the regular expressions module is called `re`:

```python
import re

```

## Queston 1: Matching words
We'll practice on the file `words.txt.gz` included with the problem set.

**1(a)** (2 pts) Read the file, so that you have a list of strings, one for each word. Save this list as `words`.

```python
import gzip
words = []
with gzip.open('words.txt.gz', 'r') as f1:
    con = f1.readlines()
    for line in con:
        words.append(''.join(re.findall(r'\w+', line.decode('utf-8'))))

```

```python
grader.check("q1a")

```

Output:
```
q1a results: All test cases passed!
```

**1(b)** (2 pts) Write a regular expression that matches any three consecutive vowels, and use it to produce a list named `three_vowels` of all such words that occur in this data set.

```python
vowels = '(a|o|e|i|u){3}'
three_vowels = []
for element in words:
    if re.match(vowels, element):
        three_vowels.append(element)
    else:
        continue
three_vowels

```

Output:
```
['aeolian',
 'aeon',
 'aeonian',
 'aeonic',
 'aeons',
 'aoudad',
 'aoudads',
 'eau',
 'eaux',
 'euonymus',
 'euonymuses',
 'oeillade',
 'oeillades',
 'oeuvre',
 'oeuvres',
 'ouabain',
 'ouabains',
 'ouistiti',
 'ouistitis']
```

```python
grader.check("q1b")

```

Output:
```
q1b results: All test cases passed!
```

**1(c)** (2 pts) Write a regular expression that a single word containing an even number (greater than `0`) of the letter `e`, and use it to produce a list named `even_e` of all such words in the data set.

```python
even_e = []
for element in words:
    if re.findall('e', element):
        if len(re.findall('e', element)) % 2 == 0:
            even_e.append(element)

```

```python
re.findall('e', "eeaeeaae")

```

Output:
```
['e', 'e', 'e', 'e', 'e']
```

```python
grader.check("q1c")

```

Output:
```
q1c results: All test cases passed!
```

**1(d)** (2 pts) Write a regular expression that matches any string that begins and ends with a consonant, and has no consonants in between. Your answer should be a string named `consonants_begin_end`.

```python
consonants_begin_end = '^[^aeoui]([aeiou]|\W)*[^aeoui]$'

```

```python
grader.check("q1d")

```

Output:
```
q1d results: All test cases passed!
```

**1(e)** (2 pts) Write a regular expression that matches any word *at least four letters long* whose last two letters are the first two letters in reverse order, and use it to produce a list named `fwd_2_rev` of all such words in the dataset. (An example of such a word is `cardiac`.)

```python
fwd_2_rev = []
for element in words:
    if re.match(r'^(.)(.).*\2\1$', element):
        if len(element) >= 4:
            fwd_2_rev.append(element)

```

```python
grader.check("q1e")

```

Output:
```
q1e results: All test cases passed!
```

## Question 2: Filtering Internet traffic
In this problem, you'll get a taste of a more realistic application of
regular expressions. The file `SkypeIRC.txt.gz` contains a capture of Internet traffic generated
by a laptop while it was running the program Skype. Each line represents a single "packet" of data that was either sent or received by the machine.

The first line of the file is:
```
    1   0.000000  192.168.1.2 → 212.204.214.114 IRC 96 Request (ISON)
```
The first two fields are a counter and timestamp indicating the total number of packets captured so far, and the time (in seconds) that has elapsed since the capture was initiated. Following that, `192.168.1.2 → 212.204.214.114` indicates that the packet was being sent from the [IP address](https://en.wikipedia.org/wiki/IP_address) `192.168.1.2` to `212.204.214.114`. The remainder of the line is a description of the packet contents, which varies from packet to packet.

**2(a)** (2 pts) Construct a regular expression named `zero_to_255` that matches a single line containing a between 0 and 255, inclusive. Numbers with leading zeros should not match (except for 0 itself), so "40" should match, but "040" should not. We will surround your regular expression with the anchors "^" and "$" when testing; you do not need to include them.

```python
# contents=[]
# with gzip.open('SkypeIRC.txt.gz', 'r') as f1:
#     con = f1.readlines()
#     for line in con:
#         contents.append(''.join(re.findall(r'.+',line.decode('utf-8'))))
zero_to_255 = '(1[0-9][0-9]|2[0-4][0-9]|25[0-5]|[1-9]?[0-9])'

```

```python
contents = []
with gzip.open('SkypeIRC.txt.gz', 'r') as f1:
    con = f1.readlines()
    for line in con:
        contents.append(''.join(re.findall(r'.+', line.decode('utf-8'))))
contents

```

Output:
```
['    1   0.000000  192.168.1.2 → 212.204.214.114 IRC 96 Request (ISON)',
 '    2   0.125852 212.204.214.114 → 192.168.1.2  TCP 66 6667 → 2848 [ACK] Seq=1 Ack=31 Win=57890 Len=0 TSval=2196036564 TSecr=14215752',
 '    3   0.137361 212.204.214.114 → 192.168.1.2  IRC 112 Response (303)',
 '    4   0.137413  192.168.1.2 → 212.204.214.114 TCP 66 2848 → 6667 [ACK] Seq=31 Ack=47 Win=8011 Len=0 TSval=14215889 TSecr=2196036565',
 '    5   0.235960  192.168.1.2 → 192.168.1.1  DNS 84 Standard query 0x311f PTR 2.1.168.192.in-addr.arpa',
 '    6   0.236116  192.168.1.2 → 192.168.1.1  DNS 88 Standard query 0x3120 PTR 114.214.204.212.in-addr.arpa',
 '    7   0.270252  192.168.1.1 → 192.168.1.2  DNS 84 Standard query response 0x311f PTR 2.1.168.192.in-addr.arpa',
 '    8   0.294105  192.168.1.1 → 192.168.1.2  DNS 123 Standard query response 0x3120 PTR 114.214.204.212.in-addr.arpa PTR sterling.freenode.net',
 '    9   0.985963  192.168.1.2 → 192.168.1.1  DNS 81 Standard query 0x3121 A sterling.freenod
```

```python
grader.check("q2a")

```

Output:
```
q2a results: All test cases passed!
```

**2(b)** (2 pts) Using your solution in the previous step, build a regular expression named `ip` that matches IP addresses. For the purpose of this exercise, we will define an IP address to be a string of the form `A.B.C.D`, where $A$, $B$, $C$, and $D$ are numbers between 0 and 255.

```python
ip = zero_to_255+'\.'+zero_to_255+'\.'+zero_to_255+'\.'+zero_to_255

```

```python
grader.check("q2b")

```

Output:
```
q2b results: All test cases passed!
```

**2(c)** (2 pts) Create a regular expression that extracts the sender and receiever IP addresses from each line of the SkypeIRC log file, and use a `collections.Counter` to count up all the unique sender and receiver addresses in the data. What do you think is the IP address of the laptop on which the capture was performed?

192.168.1.2

```python
from collections import Counter
senders = Counter()
receivers = Counter()
contents = []

with gzip.open('SkypeIRC.txt.gz', 'r') as f1:
    con = f1.readlines()
    for line in con:
        contents.append(''.join(re.findall(r'.+', line.decode('utf-8'))))
s_list = []
r_list = []


for element in contents:
    temp = re.findall(ip, element)
    if len(temp) >= 2:
        s1, s2, s3, s4 = temp[0]
        str_temp1 = s1 + "." + s2 + "." + s3 + "." + s4
        s_list.append(str_temp1)
        r1, r2, r3, r4 = temp[1]
        str_temp2 = r1 + "." + r2 + "." + r3 + "." + r4
        r_list.append(str_temp2)

senders = Counter(s_list)
receivers = Counter(r_list)
receivers

```

Output:
```
Counter({'212.204.214.114': 159,
         '192.168.1.2': 1068,
         '192.168.1.1': 359,
         '71.10.179.129': 43,
         '172.200.160.242': 41,
         '86.128.100.24': 1,
         '68.95.198.126': 13,
         '86.128.187.110': 2,
         '24.177.122.79': 27,
         '68.32.70.119': 9,
         '86.197.95.238': 2,
         '86.31.70.81': 2,
         '86.220.100.25': 2,
         '86.130.63.111': 1,
         '80.216.195.140': 1,
         '65.190.6.124': 4,
         '165.124.253.241': 2,
         '81.236.228.111': 1,
         '86.142.117.124': 1,
         '86.142.178.222': 1,
         '86.128.163.125': 1,
         '86.134.79.66': 1,
         '86.136.192.70': 1,
         '80.186.57.176': 1,
         '68.206.150.243': 29,
         '82.128.194.105': 1,
         '204.152.205.205': 4,
         '195.215.8.141': 9,
         '83.147.171.206': 1,
         '172.207.190.217': 1,
         '190.44.165.86': 5,
         '83.130.238.168': 5,
         '83.5.146.219': 1,
         '82.253.163.
```

```python
Counter([1, 2, 3])

```

Output:
```
Counter({1: 1, 2: 1, 3: 1})
```

```python
grader.check("q2c")

```

Output:
```
q2c results: All test cases passed!
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()

```

Output:
```
q1a results: All test cases passed!

q1b results: All test cases passed!

q1c results: All test cases passed!

q1d results: All test cases passed!

q1e results: All test cases passed!

q2a results: All test cases passed!

q2b results: All test cases passed!

q2c results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload this .zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)

```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1a (points: 2)
```
>>> assert len(words) == 113809
```


Tests for q1b (points: 2)
```
>>> assert len(three_vowels) == 19
```
```
>>> assert "oeillade" in three_vowels
```
```
>>> three_vowels_soln = set([w for w in words if re.match(r"[aeuio]{3}", w)])
>>> assert set(three_vowels) == set(three_vowels_soln)
```


Tests for q1c (points: 2)
```
>>> assert "eye" in even_e
>>> assert "eels" in even_e
>>> assert "storekeeper" in even_e
>>> assert "bird" not in even_e
>>> assert "fed" not in even_e
>>> assert "eleven" not in even_e
```


Tests for q1d (points: 2)
```
>>> assert isinstance(consonants_begin_end, str)
```
```
>>> regexp = re.compile(consonants_begin_end)
>>> assert regexp.match("dad")
>>> assert regexp.match("doom")
>>> assert regexp.match("mom")
>>> assert regexp.match("me ad")
>>> assert regexp.match("woooo!!!!! ah")
>>> assert not regexp.match("")
>>> assert not regexp.match("dad & mom")
>>> assert not regexp.match("dune")
```


Tests for q1e (points: 2)
```
>>> assert "cardiac" in fwd_2_rev
>>> assert "defeated" in fwd_2_rev
>>> assert "sees" in fwd_2_rev
```


Tests for q2a (points: 2)
```
>>> assert isinstance(zero_to_255, str)
>>> regexp = re.compile(r"^" + zero_to_255 + "$")
>>> assert not regexp.match("001")
>>> assert not regexp.match("123 234")
>>> for i in range(256):
...     assert regexp.match(str(i)), f"{str(i)} did not match"
...     assert regexp.match(str(i)).groups() == (str(i), ), f"{str(i)} should match exactly once"
```


Tests for q2b (points: 2)
```
>>> assert isinstance(ip, str)
>>> regexp = re.compile(r"^" + ip + r"$")
>>> assert regexp.match("127.0.0.1")
>>> assert regexp.match("172.217.2.36")
>>> assert regexp.match("141.211.243.251")
>>> assert not regexp.match("1.2.3.400")
>>> assert not regexp.match("1.2.3.4.5")
```


Tests for q2c (points: 2)
```
>>> assert senders['212.204.214.114'] == 141
```
```
>>> assert receivers['212.204.214.114'] == 159
```
